rule_files:
    - rules.yml

evaluation_interval: 1m

tests:

    - name: RDSDiskSpacePrediction
      interval: 1m
      input_series:
          - series: 'rds_free_storage_bytes{aws_account_id="111111111111",aws_region="eu-west-3",dbidentifier="db1"}'
            values: '21474836480-536870912x60'  # 10GB - 500 MB/minutes
          - series: 'rds_allocated_storage_bytes{aws_account_id="111111111111",aws_region="eu-west-3",dbidentifier="db1"}'
            values: '21474836480x60'  # 20GB
      alert_rule_test:
          - alertname: RDSDiskSpacePrediction
            eval_time: 15m
            exp_alerts:
                - exp_labels:
                      aws_account_id: 111111111111
                      aws_region: eu-west-3
                      dbidentifier: db1
                      severity: critical
                  exp_annotations:
                      description: "db1 will run out of disk space in 4 hours"
                      summary: "Predic out of disk space in 4 hours"
                      runbook_url: "https://qonto.github.io/database-monitoring-framework/0.0.0/runbooks/rds/RDSDiskSpacePrediction"
